# Seminar 14 - Combinatorics, Graph, Boolean Logic

## Introduction to Network Science

## HW 9 - Combinatorics

Please, fill in before you start:

First Name:

Last Name:

Group:

### After the assigment is done, please, push it to your [private GitHub repository](https://docs.github.com/en/github/administering-a-repository/managing-repository-settings/setting-repository-visibility) and invite [Majid-Sohrabi](https://github.com/Majid-Sohrabi)  and [Sifei-Meng](https://github.com/mengsifei) as [collaborator](https://docs.github.com/en/account-and-profile/setting-up-and-managing-your-github-user-account/managing-access-to-your-personal-repositories/inviting-collaborators-to-a-personal-repository).

Complete all task and send the notebook as HW.


- Please create a private GitHub repository and add Majid and Sifei as collaborator. Commit your solution (jupyter notebook) to your github repository.

- Deadline: **May 17, 2024, 11:59 pm**.

- Please keep the file name but replace **YOUR_NAME** part by your name: **[HW9_Combinatorics]-YOUR_NAME**

```python
# !pip install networkx
```

```python
import numpy as np
import networkx as nx
import matplotlib.pyplot as plt
import requests
from zlib import adler32
```

```python
nx.__version__
```

Output:
```
'3.3'
```

In case of an exception, install networkx 2.5 using `pip` as follows:
```
pip uninstall networkx --yes
pip install networkx==2.5
```

```python
# assert nx.__version__=='2.5'
assert nx.__version__=='3.3'
```

### Task 1. Reading a graph from a file (0.5 point)

```python
url = 'https://raw.githubusercontent.com/network-science-course/network-science-course/main/datasets/assignment_1_task_1_part_1_edges.txt'
open('assignment_1_task_1_part_1_edges.txt', 'wb').write(requests.get(url).content)

url = 'https://raw.githubusercontent.com/network-science-course/network-science-course/main/datasets/assignment_1_task_1_part_2_adjlist.txt'
open('assignment_1_task_1_part_2_adjlist.txt', 'wb').write(requests.get(url).content);

```

Usually graph is a sparse data structure. So in most cases data is stored in sparse table formats
1. edgelist
2. adjlist

Edgelist could be depicted as table with at least two columns: source and target. Row with values `{source: i, target: j}` shows that graph has an edge between nodes `i` and `j`.

Adjlist contains index of node and list of all adjacent nodes to it.

Most of other specific formats are similar except of syntax and metadata description.

For most cases we can use given functions in networkx:

* `nx.read_adjlist`
* `nx.read_edgelist`
* `nx.read_weighted_edgelist`

In the further task, read two files: `assignment_1_task_1_part_1_edges.txt`, `assignment_1_task_1_part_2_adjlist.txt` and merge it in a single graph.

<!-- grade_id: cell-3d2796d8ac544282, points: 0, grade: false, solution: true -->
```python
def read_and_merge_graphs(path1, path2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-980a87cb1b6f0dbc, points: 0.5, grade: true, solution: false -->
```python
G = read_and_merge_graphs("assignment_1_task_1_part_1_edges.txt", 
                          "assignment_1_task_1_part_2_adjlist.txt")

assert nx.number_of_nodes(G) == 504
assert nx.number_of_edges(G) == 369
```

### Task 2. Nodes labeling (1 point) --- FOR HW

Networkx read node names in text format by default, so it could be useful to preprocess nodes. For example by converting string in consecutive integer. For example it will help to work with operations over adjacency matrices and vectors of labels in DGL or PyG frameworks.

```python
list(G.nodes)[:5]
```

<!-- grade_id: cell-00a45e63001782dc, points: 0, grade: false, solution: true -->
```python
def convert_node_names_to_int(G):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8e24f330185f725a, points: 1, grade: true, solution: false -->
```python
nodes = list(convert_node_names_to_int(G).nodes())

assert sorted(nodes) == list(range(len(nodes)))
```

### Task 3. Directed graph (0.5 point)

The direction in graphs could be crucial feature for many tasks, e.g., link analysis, link prediction.

We can handle directed graphs with `nx.DiGraph` class. We can use conversion method `nx.to_directed` on the given graph or read it directly passing `nx.DiGraph` constructor to the `nx.read_*` methods.

<!-- grade_id: cell-4ff100c55840771c, points: 0, grade: false, solution: true -->
```python
def convert_to_directed_graph(G):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-678028b3a41d774a, points: 0, grade: false, solution: true -->
```python
def read_directed_graph(path1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-67792dd08b2f0149, points: 0.5, grade: true, solution: false -->
```python
assert convert_to_directed_graph(G).is_directed()
assert read_directed_graph("assignment_1_task_1_part_1_edges.txt").is_directed()
```

### Task 4. Multiple directed graph (1 point) - FOR HW

One more interesting graph case is the Heterogeneous Graph. It could have several edges between two nodes. For example, the airline network has as many edges, as the number of itineraries between two cities. All of these edges will have their own properties like departure and arrival time, prices and etc.

For such case there is a class `nx.MultiDiGraph`

In this task we will create a `nx.MultiDiGraph` with two nodes `0` and `1` and three edges between them: one from `0` to `1` and two from `1` to `0`.

<!-- grade_id: cell-66bbf015fd2e17b2, points: 0, grade: false, solution: true -->
```python
def create_graph_multi_digraph():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d33184ec5e350cf2, points: 1, grade: true, solution: false -->
```python
Gm = create_graph_multi_digraph()
assert list(Gm.edges()) == [(0, 1), (1, 0), (1, 0)]
```

### Task 5. Basic statistics (0.7 point)

Networkx provides a simple interface for getting basic graph statistics:
* `nx.radius`
* `nx.diameter`
* `nx.average_shortest_path_length`
* `nx.average_clustering`

For example, here are statistics of a path graph with 3 nodes:

```python
# Path graph - no cycle - each node connected to the next one

G = nx.path_graph(3)
print('Radius: ', nx.radius(G))     # from center to other nodes
print('Diameter: ', nx.diameter(G)) # from leaves to other leaves
print('Average shortest path length: ', nx.average_shortest_path_length(G)) # 0-->1=1, 0-->2=2, 1-->2=1 (4/3=1.33)
print('Average clustering coefficient: ', nx.average_clustering(G))  # leaf by defult 0, for node one, 0 and 2 are not connected then it's 0
nx.draw_networkx(G)
```

Your task is to create graphs by given statistics.

Write a function `graph1` that returns a graph with a radius 3 and diameter 3.

<!-- grade_id: cell-1b71ec27d5c70dc5, points: 0, grade: false, solution: true -->
```python
def graph1():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-02d925c056e1f040, points: 0.2, grade: true, solution: false -->
```python
G = graph1()
assert nx.radius(G) == nx.diameter(G) == 3
```

```python
nx.draw_networkx(G)
```

Write a function `graph2` that returns a graph with a diameter of at least 5 time greated than the average shortest path lengths.

<!-- grade_id: cell-e6b06880520f2cf7, points: 0, grade: false, solution: true -->
```python
def graph2():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a6c86b83e11199d8, points: 0.2, grade: true, solution: false -->
```python
G = graph2()
assert nx.diameter(G) >= 5 * nx.average_shortest_path_length(G)
```

```python
nx.draw_networkx(G)
```

Write a function `graph3` that returns a graph with an average clustering coefficient between 0.1 and 0.11.

<!-- grade_id: cell-eda0d295af4fc263, points: 0, grade: false, solution: true -->
```python
def graph3():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d4aeb27141847c83, points: 0.3, grade: true, solution: false -->
```python
G = graph3()
print(nx.average_clustering(G))
assert 0.1 <= nx.average_clustering(G) <= 0.11
```

```python
nx.draw_networkx(G)
```

### Task 6. Graph layout (0.7 points) - FOR HW

A graph layout is a dictionary with a location of every node: keys are nodes and values are x, y positions. By default in `nx.draw_networkx`, it is the spring layout — a force-directed representation of the network treating edges as springs holding nodes close, while treating nodes as repelling objects.

```python
G = nx.barabasi_albert_graph(10, 2)
layout = nx.spring_layout(G)
layout
```

```python
nx.draw_networkx(G, pos=layout)
```

The Kamada Kawai layout is based on path-length cost-function. Sometimes, it is quite similar to the spring layout, however a significant difference is that the spring layout is random and the Kamada Kawai is deterministic.

```python
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx(G, pos=layout)
```

Circular layout allows to draw a graph where nodes are in a circle.

```python
layout = nx.circular_layout(G)
plt.figure(figsize=(5, 5))
nx.draw_networkx(G, pos=layout)
```

Bipartite layout allows to draw bipartite graph.

```python
B = nx.Graph()
B.add_nodes_from([1, 2, 3, 4], bipartite=0)
B.add_nodes_from(["a", "b", "c"], bipartite=1)
B.add_edges_from([(1, "a"), (1, "b"), (2, "b"), (2, "c"), (3, "c"), (4, "a")])
layout = nx.bipartite_layout(B, [1, 2, 3, 4])
nx.draw_networkx(B, pos=layout)
```

There are other layouts in networkx, also there are some useful parameters that you can find in documentation. Next, let us visualize a coappearance network of characters in The Matrix — a 1999 American science fiction action film written and directed by the Wachowskis. The dataset is taken from https://moviegalaxies.com/.

```python
url = 'https://raw.githubusercontent.com/network-science-course/network-science-course/main/datasets/the_matrix_graph.gexf'
open('the_matrix_graph.gexf', 'wb').write(requests.get(url).content);
```

```python
def get_matrix_graph():
    G = nx.read_gexf('the_matrix_graph.gexf')
    mapping = {}
    for node in G.nodes:
        mapping[node] = G.nodes[node]['label']
    G = nx.relabel_nodes(G, mapping)
    return G
matrixG = get_matrix_graph()
plt.figure(figsize=(13, 7))
layout = nx.kamada_kawai_layout(matrixG)
nx.draw_networkx(matrixG, pos=layout)
```

Your task is to draw the matrix graph using the kamada kawai layout with nonlinear modification to obtain a denser layout away from the NEO node. Let us define a proximity coefficient using Gaussian kernel

$$\text{coef}(i) = \exp\left(-\frac{\text{distance}^2}{2\sigma^2}\right)$$

where a _distance_ is the Euclidean distance from the node $i$ to the NEO node, $\sigma$ is the constant. Then the nodes should be moved towards the NEO node such that

$$\text{distance} = \text{coef}(i) \times \text{distance}$$

The simplest way to do it is to center the layout so that the NEO node is in x=0 y=0 and then to multiply all node positions by the proximity coefficient.

Write a function `dense_neo_layout` that takes a layout, sigma and returns a new (dense) layout.

_Hint: use `np.linalg.norm` to calculate Euclidean norm._

<!-- grade_id: cell-6dc95d64261dffa6, points: 0, grade: false, solution: true -->
```python
def dense_neo_layout(layout, sigma):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fea67b2cc10bdbd6, points: 0.7, grade: true, solution: false -->
```python
layout = nx.kamada_kawai_layout(matrixG)
layout = dense_neo_layout(layout, sigma=1.5)
assert type(layout) == dict
assert np.all(layout['NEO'] == [0, 0])
assert np.all(layout['REX'].round(2) == [-0.67,  0.58])
assert adler32(str(layout['AGENT BROWN'].round(2)[0]).encode()) == 32309448
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(matrixG, layout)
```

### Task 7. Nodes visualization (1 point)  - FOR HW

We can set a node size using `draw_networkx` with a parameter `node_size`. It should be a list of values ordered by `G.nodes`. For example, consider a simple path graph.

```python
G = nx.path_graph(3)
G.nodes
```

Let us set a size 200 for the node 0, 300 for 1, 400 for 2.

```python
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx(G, pos=layout, node_size=[200, 300, 400])
```

To color nodes, we can use `node_color`. It should be a list of colors:
* a tuple of float RGB 0-1 values: (red, green, blue)
* a tuple of float RGBA 0-1 values: (red, green, blue, alpha)
* a string — color name

```python
nx.draw_networkx(
    G,
    pos=layout,
    node_color=[(0.4, 0.8, 0, 0.6), (0.7, 0.4, 1, 0.6), (1, 0, 0.5, 0.6)])
```

```python
nx.draw_networkx(
    G,
    pos=layout,
    node_color=['green', 'cyan', 'pink'])
```

Also we can use a colors from matplotlib colormaps. Discrete:

```python
colors = plt.cm.Set1.colors
nx.draw_networkx(
    G,
    pos=layout,
    node_color=[colors[0], colors[1], colors[2]])
```

Or continuous:

```python
rainbow = plt.cm.rainbow
nx.draw_networkx(
    G,
    pos=layout,
    node_color=[rainbow(0.2), rainbow(0.8), rainbow(0.9)])
```

Another option of coloring nodes is to pass a list of single values by the `node_color` and colormap by the `cmap` parameter. In this case, the lowest value will assigned to the first color in colormap and vise versa.

```python
nx.draw_networkx(
    G,
    pos=layout,
    cmap = plt.cm.hot,
    node_color=[1, 2, 3])
```

We can forcibly set a color range using `vmin` and `vmax`.

```python
nx.draw_networkx(
    G,
    pos=layout,
    cmap = plt.cm.hot,
    node_color=[1, 2, 3],
    vmin=0, vmax=4)
```

We can border nodes using `edgecolors` and `linewidths` parameters.

```python
nx.draw_networkx(
    G=G, 
    pos=layout,
    cmap = plt.cm.hot, 
    node_color=[1, 2, 3],
    vmin=0, vmax=4, 
    edgecolors='black', 
    linewidths=1)
```

Your task is to adjust a size and color of nodes in the matrix graph, size and color are taken from a node degree and node attribute `G.nodes[node]`.

```python
matrixG.nodes['NEO']
```

Write a function `drawing_nodes_params` that takes a graph and returns a dictionary with parameters for `nx.draw_networkx` of the form `{'G': graph, 'node_color': color, ...}`. Parameters should contain a graph, kamada kawai layout, node colors from `viz.color` attribute. The colors should be converted from RGB range 0-255 to RGB range 0-1. The node size is a node degree `G.degree[node]` multiplied by 100.

<!-- grade_id: cell-23d41af4fdb48d3c, points: 0, grade: false, solution: true -->
```python
def drawing_nodes_params(graph):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ce378d189964cb4f, points: 0.5, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
params = drawing_nodes_params(matrixG)
assert type(params) == dict
keys = ['G', 'pos', 'node_color', 'node_size']
assert len(set(params.keys()).intersection(keys)) == 4
assert np.all(params['node_color'][10].round(2) == [0.37, 0.79, 0.75])
assert adler32(str(params['node_color'][0][1].round(2)).encode()) == 19071123
assert params['node_size'][0] == 200
assert params['node_size'][1] == 600
assert nx.to_numpy_array(params['G']).sum() == 326.0
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(**params )
```

Sometimes, it will be useful to draw nodes in nonlinear scale. If we want to make node sizes more comparable, we can use polynomial or exponential scale. On the other hand, we can use logarithmic scale to make sizes more similar.

Write a function `drawing_nodes_polyparams` that returns the same dictionary as the `drawing_nodes_params` function, but the node size is a sum of node degree in a power 3 and 10.

$$\text{size} = \text{degree}^3 + 10$$

<!-- grade_id: cell-7056cc638b9de958, points: 0, grade: false, solution: true -->
```python
def drawing_nodes_polyparams(graph):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6b4ea9133d917597, points: 0.5, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
params = drawing_nodes_polyparams(matrixG)
assert type(params) == dict
keys = ['G', 'pos', 'node_color', 'node_size']
assert len(set(params.keys()).intersection(keys)) == 4
assert np.all(params['node_color'][10].round(2) == [0.37, 0.79, 0.75])
assert adler32(str(params['node_color'][0][1].round(2)).encode()) == 19071123
assert params['node_size'][0] == 18
assert params['node_size'][1] == 226
assert nx.to_numpy_array(params['G']).sum() == 326.0
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(**params )
```

### Task 8. Edges visualization, weighted graphs (2 points)  - FOR HW

We can set an edge width using `draw_networkx` with a parameter `width`. It should be a list of values sorted by G.edges. For example, consider a simple path graph.

```python
G = nx.path_graph(4)
G.edges
```

Let us set a width 1 for the edge 0-1, 5 for 1-2, 10 for 2-3.

```python
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx(G, pos=layout, width=[1, 5, 10])
```

If we create a graph by a weighted ajacency matrix, weights will be stored in an edge attribute `weight`.

```python
W = np.array([[0, 1, 0, 0], 
              [1, 0, 5, 0], 
              [0, 5, 0, 10], 
              [0, 0, 10, 0]])
G = nx.from_numpy_array(W)
```

In this case, we can get a value using `G.edges[edge]` view

```python
G.edges[(0, 1)]
```

Or we can get all values using `nx.get_edge_attributes`

```python
nx.get_edge_attributes(G, 'weight')
```

```python
weights = list(nx.get_edge_attributes(G, 'weight').values())
nx.draw_networkx(G, pos=layout, width=weights)
```

If we want to apply specific color or style to some edges, it is possible by `edge_color` and `style` parameters. For example, we can draw an edge with a weight 10 as a red dotted line.

```python
nx.draw_networkx(
    G, pos=layout, 
    width=[1, 5, 10], 
    edge_color=['black', 'black', 'red'], 
    style=['solid', 'solid', 'dotted'])
```

Your task is to draw the matrix graph where an edge width is 0 if a weight is 1, otherwise it is a weight value itself. An edge color is `red` if weight is at least 5 and `black` otherwise.

<!-- grade_id: cell-2497d776bb6eb00c, points: 0, grade: false, solution: true -->
```python
def drawing_edges_params(graph):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b220499d1d3b12af, points: 2, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
params = drawing_edges_params(matrixG)
assert type(params) == dict
keys = ['G', 'pos', 'edge_color', 'width']
assert len(set(params.keys()).intersection(keys)) == 4
assert nx.to_numpy_array(params['G']).sum() == 326.0
params['width'][0] == 2.0
assert np.array(params['width']) @ params['width'] == 508
assert adler32(str(params['edge_color']).encode()) == 646700613
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(**params)
```

### Task 9. Labels visualization (0.6 point)  - FOR HW

In this task, we will modify some label parameters: font size, position, visibility. In networkx, we cannot pass a list of font sizes to `draw_networkx`, so we will use drawing by layers using other methods:
* `draw_networkx_nodes` — draw the nodes only
* `draw_networkx_edges` — draw the edges only
* `draw_networkx_labels` — draw the node labels only
* `nx.draw_networkx_edge_labels` — draw the edge labels only

Also we can separately adjust an `alpha` value for nodes, edges, labels. Here is a simple example, let us draw a path graph with 4 nodes where the node 0 is labeled as "a", 3 is "b", the edge 1-2 is labeled as "edge", and edges are transparent with `alpha=0.2`.

```python
G = nx.path_graph(4)
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx_nodes(G, pos=layout)
nx.draw_networkx_edges(G, pos=layout, alpha=0.2)
nx.draw_networkx_labels(G, pos=layout, labels={0: 'a', 3: 'b'})
nx.draw_networkx_edge_labels(G, pos=layout, edge_labels={(1, 2): 'edge'})
plt.show()
```

If we want to adjust a font size, it is possible by repeated drawing labels with different `font_size` values. For example:

```python
G = nx.path_graph(4)
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx_nodes(G, pos=layout)
nx.draw_networkx_edges(G, pos=layout, alpha=0.2)
nx.draw_networkx_labels(G, pos=layout, labels={3: 'b'}, font_size=10)
nx.draw_networkx_labels(G, pos=layout, labels={0: 'a'}, font_size=20)
nx.draw_networkx_edge_labels(G, pos=layout, edge_labels={(1, 2): 'edge'})
plt.show()
```

Your task is to draw the matrix graph with settings: 
* labels are invisible if a node degree less than 6
* a font size is a node degree
* all labels are slided down so that $y = y - 0.025 \sqrt{\text{degree}}$

Write a function `labels_list_parameters` that takes a graph, layout and returns a list of parameters for repited drawing labels using `nx.draw_networkx_labels`. Each element of the list is a dictionary that can be passed to `nx.draw_networkx_labels`. The first dictionary is for the node degree 6, the second is for 7 and so on.

```python
def labels_list_parameters(graph, layout):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-93a09c03f82c3cfb, points: 0.6, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
layout = nx.kamada_kawai_layout(matrixG)
list_params = labels_list_parameters(matrixG, layout)
assert type(list_params) == list
assert type(list_params[0]) == dict
assert len(list_params) == 8
assert nx.to_numpy_array(list_params[0]['G']).sum() == 326
assert list_params[0]['font_size'] == 6
assert adler32(str(np.array(list_params[7]['font_size'])).encode()) == 10027110
assert np.all(list_params[0]['pos']['MAN'].round(2) == [-0.16, 0.05])
assert adler32(
    str(list_params[0]['pos']['AGENT JONES'].round(2)[1]).encode()) == 46465263
assert adler32(
    str(list_params[7]['pos']['NEO'].round(2)[1]).encode()) == 46792946
assert list_params[1]['labels']['AGENT SMITH'] == 'AGENT SMITH'
assert list_params[1]['labels'].get('AGENT JONES') is None \
    or list_params[1]['labels']['AGENT JONES'] == ''
assert len([val for val in  list_params[2]['labels'].values() if val != '']) == 3
```

```python
plt.figure(figsize=(13, 7))
layout = nx.kamada_kawai_layout(matrixG)
nx.draw_networkx_nodes(matrixG, pos=layout)
nx.draw_networkx_edges(matrixG, pos=layout, alpha=0.2)
for params in labels_list_parameters(matrixG, layout):
    nx.draw_networkx_labels(**params)
plt.show()
```

### Task 10. Visible area (2 points)  - FOR HW

Drawing a large graph, it may be useful to zoom in an area near some nodes by `plt.xlim(x_min, x_max)` and `plt.ylim(y_min, y_max)` where `x_min` is a right border, `y_max` is a top border. 

```python
G = nx.path_graph(4)
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx(G, pos=layout)
plt.xlim(-0.5, 0.5)
plt.ylim(-0.5, 0.5)
plt.show()
```

Let us adjust a visible area and combine previous drawing techniques to get a more meaningful figure:
* take a node color from the `viz.color` node attribute converting 0-255 range to 0-1 range
* set a node size as $5\times\text{degree}^2$
* take an edge width from the `weight` edge attribute
* set an edge alpha is 0.2
* labels are invisible if a node degree less than 6
* set a font size as $3 \times \text{degree}^{2/3}$
* set a visible area such that a lower border is NEO, higher is AGENT SMITH, left is MORHEUS, right is TANK.

Write a function `get_all_params` that takes a graph and layout and returns a dictionary with keys:
* `nodes_params` — dict, parameters for `nx.draw_networkx_nodes`
* `edges_params` — dict, parameters for `nx.draw_networkx_edges`
* `labels_param_list` — list of dict, parameters for `nx.draw_networkx_edges`, for degrees 6, 7 and so on
* `area` — list of float vlues: a left, right, lower, higher borders of the visible area

<!-- grade_id: cell-563f657325a0fe0e, points: 0, grade: false, solution: true -->
```python
def get_all_params(graph, layout):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3e36285c3c4a583a, points: 2, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
layout = nx.kamada_kawai_layout(matrixG)
all_params = get_all_params(matrixG, layout)
assert type(all_params) == dict
keys = ['nodes_params', 'edges_params', 'labels_param_list', 'area']
assert len(set(all_params.keys()).intersection(keys)) == 4
assert nx.to_numpy_array(all_params['nodes_params']['G']).sum() == 326
assert np.all(all_params['nodes_params']['pos']['NEO'].round(2) == [0., -0.3])
assert np.all(all_params['nodes_params']['node_color'][4].round(2) == [1., 0.86, 0.4])
assert all_params['nodes_params']['node_size'][10] == 20
assert nx.to_numpy_array(all_params['edges_params']['G']).sum() == 326
assert np.all(all_params['edges_params']['pos']['NEO'].round(2) == [0. , -0.3])
assert all_params['edges_params']['alpha'] == 0.2
assert all_params['edges_params']['width'][20] == 2
assert type(all_params['labels_param_list']) == list
assert nx.to_numpy_array(all_params['labels_param_list'][0]['G']).sum() == 326
assert np.all(all_params['labels_param_list'][5]['pos']['NEO'].round(2) == [0. , -0.3])
assert all_params['labels_param_list'][5]['labels'].get('WOMAN') is None \
    or all_params['labels_param_list'][5]['labels']['WOMAN'] == ''
assert round(all_params['labels_param_list'][2]['font_size'], 4) == 12
assert all_params['area'][0] < all_params['area'][1]
assert all_params['area'][2] < all_params['area'][3]
```

```python
plt.figure(figsize=(10, 7))
nx.draw_networkx_nodes(**all_params['nodes_params'])
nx.draw_networkx_edges(**all_params['edges_params'])
for params in all_params['labels_param_list']:
    nx.draw_networkx_labels(**params)
plt.xlim(all_params['area'][0] - 0.15, all_params['area'][1] + 0.1)
plt.ylim(all_params['area'][2] - 0.10, all_params['area'][3] + 0.1)
plt.show()
```

```python

```